Make a variable `a` with value `"hello"`

```python
a = "hello"
```

```python
print(a)
```

Output:
```
hello
```

### web reading test

```python
import urllib
f = urllib.request.urlopen('https://data.nasa.gov/api/views/gh4g-9sfh/rows.csv')
text = f.read().decode('utf-8')
f.close()
print(text[:25])
```

Output:
```
name,id,nametype,recclass
```

```python

```

Tests for q1 (points: None)
```
>>> assert a == "hello"
```
